# Lab06: Web Scraping

```python
# import library
import requests
from bs4 import BeautifulSoup
import types, importlib, requests, sys, re

import otter
grader = otter.Notebook()
```

# Q1: Scrape Title

**Objective:**
Scrape all the titles of the books listed on the Books to Scrape website (http://books.toscrape.com/).

**Instructions:**
- Use requests to send a GET request to the Books to Scrape homepage.
- Parse the HTML content using BeautifulSoup.
- Extract the book titles, which are inside "h3" tags with a nested "a" tag, and print them.

```python
url = "https://books.toscrape.com/"

# Send GET request to the website
response = ...
# Parse the page content using BeautifulSoup
soup = BeautifulSoup(...)
# Find all the title tags and print them
titles = ...
title_list = [title.text.strip() for title in titles]
title_list
```

```python
grader.check("q1")
```

# Q2: Web Scraping (headlines)
**Objective:**
Scrape headlines from the website https://news.ycombinator.com/.

**Task:**
- Send a GET request to the URL using requests.
- Parse the HTML content with BeautifulSoup.
- Extract and print the headlines inside <a> tags with class storylink.
Handle cases where no headlines are found.

```python
# URL of the news site to scrape
url = "https://news.ycombinator.com/"

# Send GET request to the website
response = ...

# Check if the request was successful (status code 200)
if response.status_code == 200:
    print("Request successful!")
else:
    print(f"Request failed with status code: {response.status_code}")
    
# Parse the page content using BeautifulSoup
soup = ...

# Debugging: Print the raw HTML content to see what we are working with
print(soup.prettify()[:1000])  # Print first 1000 characters to inspect the HTML

# Find all the headline links (they are in <a> tags with class 'storylink')
headlines = ...
# Print each headline
title_list = [title.text.strip() for title in headlines]
title_list
```

```python
grader.check("q2")
```

# Q3. Web Scraping (links)
**Objective:**
Scrape all the links (from <a> tags) from the Wikipedia homepage (https://www.wikipedia.org/).

**Instructions:**
- Use requests to send a GET request to the Wikipedia homepage.
- Parse the HTML content using BeautifulSoup.
- Extract all the links (href attributes) from <a> tags and print them.

```python
url = "https://www.wikipedia.org/"

# Send GET request to the website
response = ...

# Parse the page content using BeautifulSoup
soup = ...

# Find all the anchor tags (<a>) which contain links
links = ...
# Print every non-empty href
href_list = [link.get("href") for link in links if link.get("href")]

```

```python
grader.check("q3")
```

# Q4. Web Scraping (quote texts)
**Objective:**
Scrape all the quote texts from the Quotes to Scrape website (https://quotes.toscrape.com/).

**Instructions:**
- Use requests to send a GET request to the Quotes to Scrape homepage.
- Parse the HTML content using BeautifulSoup.
- Extract all the quote texts, which are inside <span> tags with the class 'text', and print them.

```python
# URL of the blog to scrape
url = "https://quotes.toscrape.com/"

# Send GET request to the website
response = ...

# Parse the page content using BeautifulSoup
soup = ..

# Find all the article titles (here, they are in <span> tags with class 'text')
titles = ...
# Print each quote
title_list = [title.text.strip() for title in titles]
title_list
```

```python
grader.check("q4")
```

# Q5. Scraping a Simple Book Store Website for:
    - Book titles
    - Prices
    - Ratings
    - Availability
- URL: http://books.toscrape.com/
- Here, you will work with a scraper example (Book Titles + Prices)

```python
url = "http://books.toscrape.com/"

response = requests.get(url)

if response.status_code == 200:
    print("Request successful!")
else:
    print(f"Request failed with status code: {response.status_code}")

soup = ...

# Find all book containers
books = ...

# Extract title and price
book_data = []
for book in books:
    title = ...
    price = ...
    book_data.append((title, price))

# Show first 5
for title, price in book_data[:5]:
    print(f"{title} — {price}")
```

```python
grader.check("q5")
```

# Q6: Web Scraping eBay for Product Listings

**Objective:** Write a Python script to scrape eBay search results for a specific product (e.g., laptops), extract product details, and display them.

<!-- BEGIN QUESTION -->

**Requirements:**

- Use the requests library to send an HTTP GET request to eBay’s search results page for a given product (e.g., "laptop").
- Use the BeautifulSoup library to parse the HTML content of the response.
- Extract the following details for each product listing:
  - Title: The name or description of the product.
  - Price: The price of the product.
  - Hit Count: The number of views or watchers for the product.
  - Display the extracted data in a readable format with product index, title, price, and hit count.
  - Handle cases where some data may be missing by providing default values.

```python
url = "https://www.ebay.com/sch/i.html?_nkw=laptop"

# YOUR CODE HERE
```

<!-- END QUESTION -->

## Congratulations! You have finished Lab6! 

### Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: [1, 1, 1, 2])
```
>>> 'Himalayas' in title_list[-1]
True
```
```
>>> len(title_list) == 20
True
```
```
>>> all((isinstance(t, str) for t in title_list))
True
```
```
>>> sum(('...' in t for t in title_list)) >= 5
True
```


Tests for q2 (points: [3, 3, 3])
```
>>> title_list[-1] == 'More'
True
```
```
>>> len(title_list) == 61
True
```
```
>>> title_list[0] == '1.'
True
```


Tests for q3 (points: [3, 3, 3, 3])
```
>>> len(href_list) == 371
True
```
```
>>> expected_slice = ['//ja.wikipedia.org/', '//ru.wikipedia.org/', '//de.wikipedia.org/', '//es.wikipedia.org/']
>>> href_list[1:5] == expected_slice
True
```
```
>>> len(href_list) > 0
True
```
```
>>> all((isinstance(href, str) and href.strip() != '' for href in href_list))
True
```


Tests for q4 (points: [3, 3, 3])
```
>>> len(title_list) == 10
True
```
```
>>> all((isinstance(t, str) and t.strip() != '' for t in title_list))
True
```
```
>>> all((quote.startswith('“') and quote.endswith('”') for quote in title_list))
True
```


Tests for q5 (points: [2, 2, 2, 2])
```
>>> len(book_data) == 20
True
```
```
>>> isinstance(book_data, list)
True
```
```
>>> all((isinstance(item, tuple) and len(item) == 2 for item in book_data))
True
```
```
>>> all(('<' not in title and '>' not in title for (title, _) in book_data))
True
```
